## **Write your answers directly in the cells on this notebook.**

Use as name of the variables the ones provided as placeholders (often `sol=0`. Delete the `raise NotImplementedError()` line.

### Show the entire process, from the initial fractions to the final result. 
The probability results must be a decimal value between 0.0 and 1.0. Round your number to the third decimal (e.g 0.574).**

Consider the following boxes (two boxes (red/blue) with two types of balls (green/orange)):
![image.png](image.png)

#### Exercise 1: which is the probability of selecting the red box?

<!-- grade_id: cell-7969bda5a10dc7ac, points: 0, grade: false, solution: true -->
```python
sol = round(1/2,3)
sol
```

Output:
```
0.5
```

<!-- grade_id: cell-2b9f005060121b41, points: 1, grade: true, solution: false -->
```python
assert sol == 0.5
```

#### Exercise 2: if you want to run out of balls both boxes at the same time, how often should you select the blue box?

<!-- grade_id: cell-10a736364bdf6ef5, points: 0, grade: false, solution: true -->
```python
sol = 0
sol = round(4/12,3)

```

<!-- grade_id: cell-d06dc523a3f300ce, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.3333333333333333) < 0.01
```

#### Exercise 3: which is the probability of picking up a green ball if you are using the red box?

<!-- grade_id: cell-335a720d5716ad4b, points: 0, grade: false, solution: true -->
```python
sol = 0
sol = round(2/8,3)
```

<!-- grade_id: cell-3cc76e304b24df43, points: 1, grade: true, solution: false -->
```python
assert sol == 0.25
```

#### Exercise 4: which is the probability of picking up an orange ball (any of the boxes)?

<!-- grade_id: cell-6be260afb8041372, points: 0, grade: false, solution: true -->
```python
sol = 0
sol = round(7/12,3)

```

<!-- grade_id: cell-57ee289d4e5df3f0, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.5833333333333334) < 0.01
```

#### Exercise 5: given that you picked up a green ball, which is the probability of having selected the red box?

<!-- grade_id: cell-9c4af4d3ec207bb2, points: 0, grade: false, solution: true -->
```python
sol = 0
sol = round(2/5,3)

```

<!-- grade_id: cell-10164afff838da6f, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.125) < 0.3
```

#### Optional Exercise 6: I have chosen a ball from a box and I picked up a green one. Which is the probability of having used the blue box?

<!-- grade_id: cell-ac2e94b60b1ed6fc, points: 0, grade: false, solution: true -->
```python
# No need to use the Bayes Theorem as it has not been explained in class
sol = 0
sol = round(3/5,3)

```

<!-- grade_id: cell-1ca0be55ea9d081c, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.375) < 0.9
```

Now consider a standard deck of cards (the one used in Poker):
![deck.png](deck.png)


#### Exercise 7: which is the probability of drawing an Ace?

<!-- grade_id: cell-f7a947cdf92e17c3, points: 0, grade: false, solution: true -->
```python
# Create function that returns probability percent rounded to three decimal places

def event_probability(event_outcomes, sample_space):
    probability = 0
    return round(event_outcomes/sample_space,3)
sol = event_probability(4,52)
```

<!-- grade_id: cell-adb42109bd9b7996, points: 1, grade: true, solution: false -->
```python
# Using the previously defined function: Determine the probability of drawing an Ace
assert abs(sol-0.077 ) < 0.01
```

#### Exercise 8: which is the probability of drawing a card that is a Heart?

<!-- grade_id: cell-e261cc3812382a18, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing a heart

sol = event_probability(13,52)

```

<!-- grade_id: cell-2bcb2371e498a429, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.25 ) < 0.01
```

#### Exercise 9: which is the probability of drawing a face card (such as Jacks, Queens, or Kings)?

<!-- grade_id: cell-fc4fcd733d3e1ca7, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing a face card

sol = event_probability(12,52)
```

<!-- grade_id: cell-f17c3aa2f33022be, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.231 ) < 0.01
```

#### Exercise 10: which is the probability of drawing a face card which is also a Heart?

<!-- grade_id: cell-c50a860ddb2f18f9, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing the jack, queen or king  of hearts
sol = event_probability(3,52)
```

<!-- grade_id: cell-dcd84807c13f2154, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.058 ) < 0.01
```

#### Exercise 11: which is the probability of drawing an Ace on the second draw, if the first card drawn was a King?

<!-- grade_id: cell-c953e962deacd22d, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing an Ace after drawing a King on the first draw

sol = event_probability(4,51)
```

<!-- grade_id: cell-c0fd1832c0fa3b01, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.078 ) < 0.01
```

#### Exercise 12: which is the probability of drawing an Ace on the second draw, if the first card drawn was an Ace?

<!-- grade_id: cell-899b031f866ba3b2, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing an Ace after drawing an Ace on the first draw

sol = event_probability(3,51)
```

<!-- grade_id: cell-9109c59db8e329c8, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.059) < 0.01
```

Now consider the following situation:

You are playing Poker in the [Texas Holdem variant](https://en.wikipedia.org/wiki/Texas_hold_%27em). In case you are not familiar, it is a variant of poker in which each player has two cards and there is a set of community cards to play from. 
![texas.png](texas.png)

#### Exercise 13: what is the probability that the next card drawn will be a Diamond card?

<!-- grade_id: cell-17b8aec4013e6b9c, points: 0, grade: false, solution: true -->
```python
# Sample Space
cards = 52
player_cards = 2
turn_community_cards = 4

sol = event_probability(9,cards-player_cards-turn_community_cards)
```

<!-- grade_id: cell-ebe4abfecd5e8690, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.196) < 0.01
```

Now consider the following situation:
![image.png](texas2.png)

#### Exercise 14: which is the probability that with the next card drawn there will be five cards in sequential order? 

(Notice that any Eight ( 8, 9, 10, Jack, Queen) or any King (9, 10, Jack, Queen, King) will complete the straight)

<!-- grade_id: cell-f05c41e0e869e186, points: 0, grade: false, solution: true -->
```python
# Sample Space
cards = 52
player_cards = 2
turn_community_cards = 4

sol = event_probability(8,cards-player_cards-turn_community_cards)
```

<!-- grade_id: cell-e46db60e9b2dfa4f, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.174) < 0.01
```

## Now, we keep playing Texas Holdem 
![](https://media.giphy.com/media/EQBsWLfdxHW01Y1OHK/giphy.gif)
### But looking only at the cards we could have in our hand
### (2 possible drawing events, ignore community cards)
#### Exercise 15: which is the probability of drawing a heart OR drawing a club?

<!-- grade_id: cell-5b6c82ce08d8c47f, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing a heart or a club
# YOUR CODE HERE
sol = event_probability(26,cards) + event_probability(26,cards)*event_probability(26,cards-1)
sol = 1 - event_probability(26,cards)*event_probability(25,cards-1)
```

<!-- grade_id: cell-dedacb0828533e9c, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.755) < 0.01
```

#### Exercise 16: which is the probability of drawing an ace, a king or a queen?

<!-- grade_id: cell-0b50174237f37b03, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing an ace, king, or a queen
sol = event_probability(12,cards) + event_probability(40,cards)*event_probability(12,cards-1)
sol = 1 - event_probability(40,cards)*event_probability(39,cards-1)
```

<!-- grade_id: cell-1c3bec1e6a4f0726, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.41171499999999994) < 0.01
```

#### Exercise 17: which is the probability of drawing a heart or an ace?

<!-- grade_id: cell-3ff5cdeedc739059, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing a heart or an ace
sol = event_probability(16,cards)+event_probability(36,cards)*event_probability(16,cards-1)
```

<!-- grade_id: cell-a74509b22cdaf4bd, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.525288) < 0.01
```

#### Exercise 18: which is the probability of drawing a red card or drawing a face card?

<!-- grade_id: cell-9c43e1f91aac785f, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing a red card or a face card

sol = event_probability(32,cards)+event_probability(20,cards)*event_probability(32,cards-1)
```

<!-- grade_id: cell-47f891a62515e8d4, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.856395) < 0.01
```

#### Exercise 19: which is the probability of drawing an Ace from a deck of cards, replacing it, reshuffling the deck, and drawing another Ace?

<!-- grade_id: cell-18d5038a32ffeac9, points: 0, grade: false, solution: true -->
```python
# Sample Space
cards = 52

# Outcomes
aces = 4

sol = event_probability(4,cards)*event_probability(4,cards)
```

<!-- grade_id: cell-536e3de73c88e083, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.005929) < 0.001
```

#### Exercise 20: which is the probability of being dealt two Aces (drawing one Ace after the other starting with a full deck)?

<!-- grade_id: cell-9ec7e543991e2e7c, points: 0, grade: false, solution: true -->
```python
# Sample Space first draw
cards = 52

# Outcomes first draw
aces = 4

sol = event_probability(4,cards)*event_probability(3,cards-1)
```

<!-- grade_id: cell-0651acc2d4c97ad3, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.004542999999999999) < 0.001
```

### Permutations and combinations

#### Exercise 21: How many different 5-letter arrangements are there of the letters in the word morse?

<!-- grade_id: cell-5d16f7fa95eb0955, points: 0, grade: false, solution: true -->
```python
import math
sol = math.factorial(5)
```

<!-- grade_id: cell-290944fdf9fb212b, points: 1, grade: true, solution: false -->
```python
assert sol == 120
```

#### Exercise 22: How many different seven-letter arrangements of the letters in the word HEXAGON can be made if each letter is used only once?

<!-- grade_id: cell-2b2af5a344e47a8b, points: 0, grade: false, solution: true -->
```python
sol = math.factorial(7)
```

<!-- grade_id: cell-9eed40b79a9bc42b, points: 1, grade: true, solution: false -->
```python
assert sol == 5040
```

#### Exercise 23: Two cards are drawn at random from a standard deck of cards, without replacement. Find the probability of drawing an 8 and a queen in that order.

<!-- grade_id: cell-0772318d22f07226, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sol = event_probability(4,cards)*event_probability(4,cards-1)
```

<!-- grade_id: cell-a85db4f2cf65b50d, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.006006) < 0.001
```